```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, run the previous cell to load the provided tests.

**Recommended Readings:**

- [What is Data Science?](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

**Deadline:**

This assignment is **due Wednesday, 1/24 at 5:00pm PT**. Submissions after this time will be accepted for 24 hours and will incur a 20% penalty. Any submissions later than Thursday, 1/25 at 5:00pm PT will not be accepted unless an extension has been granted. Turn it in by Tuesday, 1/23 at 5:00pm PT for 5 extra credit points. For more information check our [policies](http://data8.org/sp24/policies/) page.

**Note: This homework has hidden tests on it. That means even though tests may say 100% passed, it doesn't mean your final grade will be 100%. We will be running more hidden tests for correctness once everyone turns in the homework.**

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged. Refer to the [policies](http://data8.org/sp24/policies/#academic-honesty) page to learn more about how to learn cooperatively.

You should start early so that you have time to get help if you're stuck. Office hours are held Monday-Friday. The schedule appears on [our website](http://data8.org/sp24/officehours/).

## 1. Scary Arithmetic

<!-- BEGIN QUESTION -->

**Question 1.1** An ad for ADT Security Systems says,

> "When you go on vacation, burglars go to work [...] According to FBI statistics, over 25% of home burglaries occur between Memorial Day to Labor Day."

Do the data in the ad support the claim that burglars are more likely to go to work during the time between Memorial Day to Labor Day? Please explain your answer. **(6 Points)**

**Note:** You can assume that "over 25%" means only slightly over. Had it been much over, say closer to 30%, then the marketers would have said so.

**Note:** Memorial Day is observed on the last Monday of May and Labor Day is observed on the first Monday of September.


_Type your answer here, replacing this text._

**SOLUTION:** No. Labor Day is around 14 weeks after Memorial Day, so the period between them is a bit more than 25% of the year (14/52 ~= 26%). 25% of burglaries happening in 25% of the year does not imply a higher rate of burglary in the summer.

<!-- END QUESTION -->

## 2. Characters in Little Women


In lecture, we counted the number of times that the literary characters were named in each chapter of the classic book, [*Little Women*](https://inferentialthinking.com/chapters/01/3/1/Literary_Characters.html?highlight=little%20women). In computer science, the word "character" also refers to a letter, digit, space, or punctuation mark; any single element of a text. The following code generates a scatter plot in which each dot corresponds to a chapter of *Little Women*. The horizontal position of a dot measures the number of periods in the chapter. The vertical position measures the total number of characters.

```python
# Just run this cell.

# This cell contains code that hasn't yet been covered in the course,
# but you should be able to interpret the scatter plot it generates.

from datascience import *
from urllib.request import urlopen
import numpy as np
%matplotlib inline

little_women_url = 'https://www.inferentialthinking.com/data/little_women.txt'
chapters = urlopen(little_women_url).read().decode().split('CHAPTER ')[1:]
text = Table().with_column('Chapters', chapters)
Table().with_columns(
    'Periods',    np.char.count(chapters, '.'),
    'Characters', text.apply(len, 0)
    ).scatter(0)
```

Output:
```
File /Library/Frameworks/Python.framework/Versions/3.10/lib/python3.10/urllib/request.py:1391, in HTTPSHandler.https_open(self, req)
   1390 def https_open(self, req):
-> 1391     return self.do_open(http.client.HTTPSConnection, req,
   1392         context=self._context, check_hostname=self._check_hostname)

File /Library/Frameworks/Python.framework/Versions/3.10/lib/python3.10/urllib/request.py:1351, in AbstractHTTPHandler.do_open(self, http_class, req, **http_conn_args)
   1348         h.request(req.get_method(), req.selector, req.data, headers,
   1349                   encode_chunked=req.has_header('Transfer-encoding'))
   1350     except OSError as err: # timeout error
-> 1351         raise URLError(err)
   1352     r = h.getresponse()
   1353 except:

URLError: <urlopen error [SSL: CERTIFICATE_VERIFY_FAILED] certificate verify failed: unable to get local issuer certificate (_ssl.c:997)>
```

**Question 1.** Around how many periods are there in the chapter with the most characters? Assign either 1, 2, 3, 4, or 5 to the name `characters_q1` below. **(4 Points)**

1. 250
2. 390
3. 440
4. 32,000
5. 40,000


**Note:** If you run into a `NameError: name 'grader' is not defined` error in the autograder cell below (and in any assignment), please re-run the first cell at the very top of this notebook!


```python
characters_q1 = 2 # SOLUTION
```

```python
grader.check("q2_1")
```

The test above checks that your answers are in the correct format. **This test does not check that you answered correctly**, only that you assigned a number successfully in each multiple-choice answer cell.

**Question 2.** Which of the following chapters has the most characters per period? Assign either 1, 2, or 3 to the name `characters_q2` below. **(4 Points)**

1. The chapter with about 60 periods
2. The chapter with about 350 periods
3. The chapter with about 440 periods


```python
characters_q2 = 1 # SOLUTION
```

```python
grader.check("q2_2")
```

Again, the test above checks that your answers are in the correct format, but not that you have answered correctly.

To discover more interesting facts from this plot, check out [Section 1.3.2](https://inferentialthinking.com/chapters/01/3/2/Another_Kind_Of_Character.html) in the textbook.

## 3. Names and Assignment Statements

**Question 1.** When you run the following cell, Python produces a cryptic error message.

```python
4 = 2 + 2
```

Output:
```
Cell In [8], line 1
    4 = 2 + 2
    ^
SyntaxError: cannot assign to literal here. Maybe you meant '==' instead of '='?
```

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `names_q1` below to indicate your answer. **(4 Points)**

1. Python is smart and already knows `4 = 2 + 2`.

2. In Python, it's a rule that the `=` sign must have a variable name to its left, and `4` isn't a variable name.

3. It should be `2 + 2 = 4`.

4. I don't get an error message. This is a trick question.


```python
names_q1 = 2 # SOLUTION
```

```python
grader.check("q3_1")
```

**Question 2.** When you run the following cell, Python will produce another cryptic error message.

```python
two = 3
six = two plus two
```

Output:
```
Cell In [12], line 2
    six = two plus two
              ^
SyntaxError: invalid syntax
```

Choose the best explanation of what's wrong with the code and assign 1, 2, 3, or 4 to `names_q2` below to indicate your answer. **(4 Points)**

1. The `plus` operation only applies to numbers, not the word "two".

2. The name "two" cannot be assigned to the number 3.

3. Two plus two is four, not six.

4. The name `plus` isn't a built-in operator; instead, addition uses `+`.


```python
names_q2 = 4 # SOLUTION
```

```python
grader.check("q3_2")
```

**Question 3.** Run the following cell.

```python
x = 2
y = 3 * x
x = 4
```

What is `y` after running this cell, and why? Choose the best explanation and assign 1, 2, 3, or 4 to `names_q3` below to indicate your answer. **(4 Points)**

1. `y` is equal to 6, because the second `x = 4` has no effect since `x` was already defined.

2. `y` is equal to 6, because `x` was 2 when `y` was assigned, and 3 * 2 is 6.

3. `y` is equal to 12, because `x` is 4 and 3 * 4 is 12.

4. `y` is equal to 12, because assigning `x` to 4 will update `y` to 12 since `y` was defined in terms of `x`.


```python
names_q3 = 2 # SOLUTION
```

```python
grader.check("q3_3")
```

## 4. Differences Between Majors

Berkeley’s Office of Planning and Analysis (OPA) provides data on numerous aspects of the campus. Adapted from the OPA website, the table below displays the number of degree recipients in three majors in the 2008-2009 and 2017-2018 academic years.

| Major                              | 2008-2009    | 2017-2018   |
|------------------------------------|--------------|-------------|
| Gender and Women's Studies         |      17      |    28       |
| Linguistics                        |      49      |    67       |
| Rhetoric                           |      113     |    56       |



**Question 1.** Suppose you want to find the **biggest** absolute difference between the number of degree recipients in the two years, among the three majors.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 49 from 67) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers. **(5 Points)**


```python
biggest_change = max(abs(17 - 28), abs(49 - 67), abs(113 - 56)) # SOLUTION
biggest_change
```

Output:
```
57
```

```python
grader.check("q4_1")
```

**Question 2.** Which of the three majors had the **smallest** absolute difference? Assign `smallest_change_major` to 1, 2, or 3 where each number corresponds to the following major:

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the smallest absolute difference. **(4 Points)** 

_Hint:_ You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each major.


```python
smallest_change_major = 1 # SOLUTION
smallest_change_major
```

Output:
```
1
```

```python
grader.check("q4_2")
```

**Question 3.**  For each major, define the “relative change” to be the following: $\large{\frac{\text{absolute difference}}{\text{value in 2008-2009}} * 100}$ 

Fill in the code below such that `gws_relative_change`, `linguistics_relative_change` and `rhetoric_relative_change` are assigned to the relative changes for their respective majors. **(5 Points)**


```python
"""# BEGIN PROMPT
gws_relative_change = (abs(...) / 17) * 100
"""; # END PROMPT
gws_relative_change = (abs(17 - 28) / 17) * 100 # SOLUTION NO PROMPT
linguistics_relative_change = (abs(49 - 67) / 49) * 100 # SOLUTION
rhetoric_relative_change = (abs(113 - 56) / 113) * 100 # SOLUTION
gws_relative_change, linguistics_relative_change, rhetoric_relative_change
```

Output:
```
(64.70588235294117, 36.734693877551024, 50.442477876106196)
```

```python
grader.check("q4_3")
```

**Question 4.** Assign `biggest_rel_change_major` to 1, 2, or 3 where each number corresponds to to the following: 

1. Gender and Women's Studies  
2. Linguistics  
3. Rhetoric

Choose the number that corresponds to the major with the biggest relative change. **(4 Points)**


```python
biggest_rel_change_major = 1 # SOLUTION
biggest_rel_change_major
```

Output:
```
1
```

```python
grader.check("q4_4")
```

## 5. Nearsightedness Study

[Myopia](https://en.wikipedia.org/wiki/Myopia), or nearsightedness, results from a number of genetic and environmental factors. In 1999, Quinn et al studied the relation between myopia and ambient lighting at night (for example, from nightlights or room lights) during childhood.

<!-- BEGIN QUESTION -->

**Question 5.1.** The data were gathered by the following procedure, reported in the study. "Between January and June 1998, parents of children aged 2-16 years [...] that were seen as outpatients in a university pediatric ophthalmology clinic completed a questionnaire on the child's light exposure both at present and before the age of 2 years." Was this study observational, or was it a controlled experiment? Explain. **(5 Points)**


_Type your answer here, replacing this text._

**SOLUTION:** It was an observational study.  The researchers didn't perform any intervention.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.2.** The study found that of the children who slept with a room light on before the age of 2, 55% were myopic. Of the children who slept with a night light on before the age of 2, 34% were myopic. Of the children who slept in the dark before the age of 2, 10% were myopic. The study concluded the following: "The prevalence of myopia [...] during childhood was strongly associated with ambient light exposure during sleep at night in the first two years after birth."

Do the data support this statement? Why or why not? You may interpret "strongly" in any reasonable qualitative way. **(5 Points)**


_Type your answer here, replacing this text._

**SOLUTION:** Yes.  There is a big difference in myopia rates between the groups.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.3.** On May 13, 1999, CNN reported the results of this study under the headline, "Night light may lead to nearsightedness." Does the original study claim that night light causes nearsightedness? **(5 Points)**


_Type your answer here, replacing this text._

**SOLUTION:** No.  The study (as quoted above) claimed only an association.

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.4.** The final paragraph of the CNN report said that "several eye specialists" had pointed out that the study should have accounted for heredity.

Myopia is passed down from parents to children. Myopic parents are more likely to have myopic children, and may also be more likely to leave lights on habitually (since the parents have poor vision). In what way does the knowledge of this possible genetic link affect how we interpret the data from the study? Explain. **(5 Points)**


_Type your answer here, replacing this text._

**SOLUTION:** If myopic parents are more likely to have myopic kids *and* leave the lights on at night, then myopic kids are more likely to have lights on at night. It is then reasonable to assume that myopic parents are a potential confounding factor that the observational study did not account for. However, we can still find the observed association even if there is no causal effect of night lights on child myopia. 

<!-- END QUESTION -->

## 6. Studying the Survivors

The Reverend Henry Whitehead was skeptical of John Snow’s conclusion about the Broad Street pump. After the Broad Street cholera epidemic ended, Whitehead set about trying to prove Snow wrong.  (The history of the event is detailed [here](http://www.ncbi.nlm.nih.gov/pmc/articles/PMC1034367/pdf/medhist00183-0026.pdf).)

He realized that Snow had focused his analysis almost entirely on those who had died. Whitehead, therefore, investigated the drinking habits of people in the Broad Street area who had not died in the outbreak.

What is the main reason it was important to study this group? Assign either 1, 2, or 3 to the name `survivor_answer` below. **(4 Points)**

1. If Whitehead had found that many people had drunk water from the Broad Street pump and not caught cholera, that would have been evidence against Snow's hypothesis.

2. Survivors could provide additional information about what else could have caused the cholera, potentially unearthing another cause.

3. Through considering the survivors, Whitehead could have identified a cure for cholera.


```python
survivor_answer = 1 # SOLUTION
```

```python
grader.check("q6_1")
```

**Note:** Whitehead ended up finding further proof that the Broad Street pump played a central role in spreading the disease to the people who lived near it. Eventually, he became one of Snow’s greatest defenders.

## 7. Policies and Administrivia

This section of the homework is to ensure that you have read over the policies and frequently asked questions for the course. 

**It's important that you read through this section of the homework very carefully**. If you can get through all of this section and are sure you have all of the correct resources set up, you will be able to focus on the actual material this semester!

Reading through the [policies](http://data8.org/sp24/policies/) and the [FAQ](http://data8.org/sp24/faq/) will help you get through this section very easily. It is recommended you do this before proceeding. 

**Question 1.** You have a question regarding the grading of your assignments that has not been previously answered on Ed or the FAQ. Who do you contact? Assign `contact` to the number corresponding to the best choice below. **(4 Points)**

1. The Instructors
2. Post on Ed
3. Contact your Lab TA


```python
contact = 3 # SOLUTION
```

```python
grader.check("q7_1")
```

**Question 2.** Why are there typically 2 items listed on Gradescope for each homework assignment? Assign `grades` to the number corresponding to the best choice below. **(4 Points)**

1. There was a mistake in the grading. I should contact someone about this.
2. One assignment is for coding questions (which I will submit through a zip file), and the other is for written responses (which I will submit through a pdf).
3. Trick question.


```python
grades = 2 # SOLUTION
```

```python
grader.check("q7_2")
```

**Question 3.** Regrade deadline dates will always be posted on the same Ed post that releases the assignment grades, common mistakes, and solutions. Can you ask for parts of your assignment regraded after the regrade request window has passed? Assign `regrade` to the number corresponding to the best choice below. **(4 Points)**

1. Yes
2. No


```python
regrade = 2 # SOLUTION
```

```python
grader.check("q7_3")
```

**Question 4.** Do you have an Gradescope account? Head to [gradescope.com](http://gradescope.com) and check if you see Data 8. If you do not, please send your Lab TA an email with your email and student ID number. 

Once you have been enrolled, go to the Data 8 Gradescope course website. At the end of the url link, you should see the six-digit number 703847. Assign `gradescope` to the integer 576157 if this is the case. **(4 Points)**


```python
gradescope = 576157 #SOLUTION
```

```python
grader.check("q7_4")
```

**Question 5.** Given the following scenarios, assign `acceptable` to the corresponding number of the scenario that is permissible given the guidelines on the [policies](http://data8.org/sp24/policies/) page. **(4 Points)**

1. Jessica gets stuck on a homework assignment, so she googles a fix. She stumbles across a pdf of the solutions for the homework assignment from a previous semester's offering of Data 8. After inspecting the solution, Nicole writes her own solution and submits the assignment.

2. After getting confused by a project, Atticus asks his friend for help. His friend Kelsey helps by walking Atticus through her own logic, without showing her code, pointing out areas that are important given the context of the question. Upon hearing his friend's logic, Atticus writes his own code and completes the project.

3. Edwin (who is in a regular lab) has an extremely busy schedule, so he really wants to leave lab early by finishing it and getting checked off. His neighbor, Gamy, simply turns his computer so Edwin can see how he completed some questions. After looking at his code, Edwin finishes the lab and gets checked off.


```python
acceptable = 2 # SOLUTION
```

```python
grader.check("q7_5")
```

**Question 6.** To make sure you have read through the [policies](http://data8.org/sp24/policies/) and the [FAQ](http://data8.org/sp24/faq/) carefully, how many HW and lab drops are there? Assign `drops` to the number corresponding to the best choice below. **(4 Points)**

1. Two homework drops and three lab drops
2. Two homework drops and two lab drops
3. Only two homework drops
4. One homework drop and two lab drops


```python
drops = 2 # SOLUTION
```

```python
grader.check("q7_6")
```

**Question 7.** Does Data 8 offer alternate final exam to those with class conflicts? Assign `exams` to the number corresponding to the best choice below. **(3 Points)**

1. Yes
2. No


```python
exams = 2 # SOLUTION
```

```python
grader.check("q7_7")
```

**Question 8:** Are you actually checking Ed? Go to this semester's [Data 8 Ed](https://edstem.org/us/courses/52859/discussion/) and find an instructor posted thread with a certain secret phrase. Assign `secret` to this secret phrase in quotes (i.e. as a string). **(4 Points)**

```python
secret = "beepboop" # SOLUTION
```

```python
grader.check("q7_8")
```

## 8. Welcome Survey

**Question 1.** Please complete the welcome survey below in order to receive credit for homework 1. Keep an eye out for the secret once} you submit! **(1 Point)**

- [Spring 2024 Welcome Survey](https://docs.google.com/forms/d/e/1FAIpQLSfFNr9JCGMgF_N3tCz7gCzQP1y5Rxg8phSodVGTRwJo2Ds3hA/viewform?usp=sf_link)

Assign `survey` to the secret phrase given at the end of the welcome survey. Make sure the phrase is in quotes (i.e. is a string)!


```python
survey = "data8 is data gr8" # SOLUTION
```

```python
grader.check("q8_1")
```

You're done with Homework 1!  

**Important submission information:** Be sure to run the tests and verify that they all pass, then choose **Save Notebook** from the **File** menu, then **run the final cell** and click the link to download the zip file. Then, go to [Gradescope](https://www.gradescope.com/courses/703847) and submit the zip file to the corresponding assignment. The name of this assignment is "HW 01 Autograder". **It is your responsibility to make sure your work is saved before running the last cell.**

Once you have submitted, your Gradescope assignment should look something like the following image if you have passed all tests.

**Note:** *This is a photo of a generic Gradescope submission result, and it does not included the same test numbers as this assignment. Please check that all test cases have passed for each question.*

<img src="gradescope.png">

## Pets of Data 8

**Piper** say congrats on finishing Homework 1! We're excited for a great semester with you!

<img src="./piper.jpg" width="30%" alt="Close up picture of dog sitting on wood path overlooking ocean."/>

Pet of the week: **Piper**

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Tests for q2_1 (points: [0, 4])
```
>>> 1 <= characters_q1 <= 5
True
```
```
>>> characters_q1 == 2
True
```


Tests for q2_2 (points: [0, 4])
```
>>> 1 <= characters_q2 <= 3
True
```
```
>>> characters_q2 == 1
True
```


Tests for q3_1 (points: [0, 4])
```
>>> 1 <= names_q1 <= 4
True
```
```
>>> names_q1 == 2
True
```


Tests for q3_2 (points: [0, 4])
```
>>> 1 <= names_q2 <= 4
True
```
```
>>> names_q2 == 4
True
```


Tests for q3_3 (points: [0, 4])
```
>>> 1 <= names_q3 <= 4
True
```
```
>>> names_q3 == 2
True
```


Tests for q4_1 (points: [0, 5])
```
>>> isinstance(biggest_change, (int, float))
True
```
```
>>> biggest_change == 57
True
```


Tests for q4_2 (points: [0, 0, 4])
```
>>> isinstance(smallest_change_major, (int, float))
True
```
```
>>> 1 <= smallest_change_major <= 3
True
```
```
>>> smallest_change_major == 1
True
```


Tests for q4_3 (points: [0, 0, 0, 1, 2, 2])
```
>>> isinstance(gws_relative_change, (int, float))
True
```
```
>>> isinstance(linguistics_relative_change, (int, float))
True
```
```
>>> isinstance(rhetoric_relative_change, (int, float))
True
```
```
>>> gws_relative_change >= 64 and gws_relative_change <= 66
True
```
```
>>> linguistics_relative_change >= 36 and linguistics_relative_change <= 38
True
```
```
>>> rhetoric_relative_change >= 50 and rhetoric_relative_change <= 52
True
```


Tests for q4_4 (points: [0, 0, 4])
```
>>> isinstance(biggest_rel_change_major, (int, float))
True
```
```
>>> 1 <= biggest_rel_change_major <= 3
True
```
```
>>> biggest_rel_change_major == 1
True
```


Tests for q6_1 (points: [0, 4])
```
>>> 1 <= survivor_answer <= 3
True
```
```
>>> survivor_answer
1
```


Tests for q7_1 (points: [4])
```
>>> contact == 3
True
```


Tests for q7_2 (points: [4])
```
>>> grades == 2
True
```


Tests for q7_3 (points: [4])
```
>>> regrade == 2
True
```


Tests for q7_4 (points: [4])
```
>>> gradescope == 576157
True
```


Tests for q7_5 (points: [4])
```
>>> acceptable == 2
True
```


Tests for q7_6 (points: [4])
```
>>> drops == 2
True
```


Tests for q7_7 (points: [3])
```
>>> exams == 2
True
```


Tests for q7_8 (points: [0, 0, 4])
```
>>> type(secret) == str
True
```
```
>>> len(secret)
8
```
```
>>> secret == 'beepboop'
True
```


Tests for q8_1 (points: [1])
```
>>> survey == 'data8 is data gr8'
True
```
